# MTH793P: Advanced Machine Learning, Semester B, 2025/2026
## Final Examination

---

**Full Name:**

<!-- grade_id: hdr-fname-a, points: 0, grade: true, solution: true -->
Shahin Hussain

**Student ID:**

<!-- grade_id: hdr-sid-a, points: 0, grade: true, solution: true -->
140474758

---

### General instructions

First, enter your **Full Name** and **Student ID** in the fields above.

**Read the following information carefully before turning to the actual exercises.**

This final project assignment is scored on a scale of 100 marks and will contribute **60% of your final mark** for this module.

This project is due by **Friday, 22 May 2026, 17:00 GMT**.

Submissions received after this deadline will be treated in accordance with the College Regulations for late submissions: Up to seven days after the deadline, late submissions will incur a penalty of 5 marks per 24 hours. Any submission received more than 168 hours (7 days) after the deadline will receive 0 marks.

The lecturer will be available to answer reasonable questions about this assessment until **Friday, 15 May 2026, 17:00 GMT**. You can contact the lecturer by email ([hong.qi@qmul.ac.uk](mailto:hong.qi@qmul.ac.uk)).

You must use this Jupyter Notebook to answer all exercises. You are not allowed to remove any cells from this notebook. You can add Markdown or code cells as needed.

You must submit this Jupyter Notebook through QMPlus. Files sent by other means, including email, will not be considered.

You cannot submit any additional files apart from this Jupyter Notebook. The ZIP file you downloaded contains various auxiliary files that you will need for the exercises. The original versions of these files will be merged into your submission automatically.


For **Markdown** exercises, **replace all occurrences of** `YOUR ANSWER HERE`  **by valid Markdown** to answer the question.

For **coding** exercises,
**replace all occurrences of**
```python
### YOUR CODE HERE
raise NotImplementedError()
```
**by valid Python code** to solve the problem.
In your solution, you can **only use** built-in Python functions and **functions from modules were imported for you in this template**.
Importing other modules is not allowed and will result in withdrawal of the corresponding marks.
If your code produces an error and you cannot fix it, write a comment to show that you are aware of the problem and indicate possible causes if you can.
You can get **partial credit**, so even if you cannot solve an exercise completely, try to answer as much as you can.


Finally, before you turn this project in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart Kernel) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All (Jupyter Notebook Version 6) or Run$\rightarrow$Run All Cells (Jupyter Notebook Version 7)). Please **do not change the file name** when you upload your Notebook to QMPlus.

<div class="alert alert-block alert-warning">
    Please answer all the questions below by writing your own Python code. Using or copy-pasting pre-existing code from any source is not permitted, and will be considered as <b>plagiarism</b> and addressed according to College Regulations. You are not allowed to collaborate with other students or ask for help from any other source, including, but not limited to, online forums or ChatGPT (or any other AI agent). By submitting your Notebook via QMPlus, you confirm that you have followed these rules and that the submitted code is your own work. Upon violation, you will fail this test with zero marks.</div>

---

### Coding style

#### Comments

Comments are an essential part of programming.
Use appropriate comments to explain the reasoning behind your code.

#### Variables & Functions

Use the variable and function names specified (**<font color=red>in red</font>**) in each exercise when writing your code. If you use other names, the testing of your code might fail, resulting in a reduced mark. For any additional variables/functions you define, choose meaningful, concise names.


It is perfectly fine if you define extra auxiliary functions to carry out a task. This is to say, even if the exercise only asks you to write a function `example_function`, you are free to add other functions like `example_function_aux`, `some_other_function`, ... and use them in your implementation of `example_function`.
Choose meaningful names for all auxiliary functions, too.

---

### Assessment criteria

The first thing we will do to evaluate your submission is to clear all outputs and run the entire notebook from the beginning to the end. It is your responsibility to ensure that the code produces **no errors**. We will not try to fix any errors when assessing your work as you have several weeks to develop this project. Instead, you will receive **zero marks for any exercises where the code cell raises an error**. If you run into an error and do not know how to fix it, comment out the critical line of code and any subsequent ones that depend on it. Add a remark indicating that you are aware of the problem and potential reasons.

If you do not know how to implement a certain operation in Python,
add a comment and explain in plain English (and in sufficient detail) what you would like to do.
You may receive partial credit for expedient ideas formulated in this way.

**10 marks (10% of the project mark)** will be awarded for **coding style**. This includes the use of appropriate variable names, suitable comments to explain complex sequences of operations, and general organization.
This is to say, even if your code is fully functional and solves the exercise, you may lose marks if the reasoning is not explained properly, or the code is not well-organized.

<div class="alert alert-block alert-danger">

**Running time:** You will be asked to code and use several machine learning methods that we studied in the lectures. If not implemented correctly, or used with the wrong parameters, these can take a long time to run. When you submit your notebook, you should make sure that:
* No single box in this notebook takes **longer than 30 seconds** to run on your personal computer.
* The entire notebook runs from beginning to end in **less than 3 minutes** on your personal computer.

You should make sure that your submitted notebook follows these limitations. Failing to do so will amount to a penalty of **20 marks**. 

---

### Admitted imports

You may only use built-in Python and the packages imported in the cell below. **Do not import additional modules.**

```python
import numpy as np
from numpy.linalg import svd, norm

import matplotlib.pyplot as plt

from skimage import data
from skimage import img_as_float
from skimage.color import rgb2gray

from sklearn.cluster import KMeans, SpectralClustering
from sklearn.preprocessing import StandardScaler

from numpy.testing import assert_equal, assert_array_equal
from numpy.testing import assert_almost_equal, assert_array_almost_equal
```

---

# Part I - Satellite Image Segmentation Using Spectral Clustering (20 marks)

In this part we use **Spectral Clustering** to segment a satellite image of land cover into
distinct regions — forest, water, urban area, and farmland — based on pixel color similarity.

The file **`satellite.npy`** contains a small vague satellite image of shape **(H, W, 3)**, with
RGB values normalized to $[0, 1]$. Each pixel is treated as a data point with a
3-dimensional color feature vector.

For this part you may use **SpectralClustering** from `sklearn.cluster` as
well as `numpy` for the manual steps.

---

The cell below loads the image into **`satellite`** and displays it.

```python
satellite = np.load('satellite_50x70.npy')
H, W = satellite.shape[:2]
print('Image shape:', satellite.shape)
plt.imshow(satellite); plt.axis('off'); plt.title('Satellite Image'); plt.show()
```

**Step 1: Prepare the pixel feature matrix.**

Reshape the satellite image into **<font color='red'>pixels</font>** — an array of shape
`(H*W, 3)` where each row is the RGB vector of one pixel. Store the image dimensions in
**<font color='red'>H</font>** and **<font color='red'>W</font>**.

<!-- grade_id: p1-code-reshape, points: 0, grade: false, solution: true -->
```python
pixels = satellite.reshape(-1, 3)
print(f"H = {H}")
print(f"W = {W}")
print(f"Pixels matrix shape: {pixels.shape}")
```

<!-- grade_id: p1-grade-reshape, points: 3, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

assert pixels.ndim == 2
assert pixels.shape[1] == 3
assert pixels.shape[0] == H * W
assert_almost_equal(pixels.min(), 0.0, decimal=3)
assert pixels.max() <= 1.0 + 1e-6
print('--- TEST PASSED ---')
```

**Step 2: Apply Spectral Clustering.**

Apply **SpectralClustering** to **`pixels`** with:
- `n_clusters = 3`
- `affinity = 'rbf'`
- `gamma = 50.0`
- `random_state = 0`
- `n_init = 5`

Store the integer labels (0–3) in **<font color='red'>labels_spectral</font>**, then reshape
to **(H, W)** and store in **<font color='red'>label_image</font>**.

Display `label_image` with `cmap='tab10'` alongside the original image.

<!-- grade_id: p1-code-spectral, points: 0, grade: false, solution: true -->
```python
# When defining spectral clustering models, we use the convention "spectral_n" -> n clusters
# Define our baseline model
spectral_3 = SpectralClustering(
    n_clusters=3,
    affinity='rbf',
    gamma=50.0,
    random_state=0,
    n_init=5
)

# Define auxillary plot function as we will use be plotting several times
def plot_clusters(spectral_cluster_model):
    clusters = spectral_cluster_model.n_clusters

    # Fit imported model to feature matrix (pixels):
    labels_spectral = spectral_cluster_model.fit_predict(pixels)

    # Reshape to original dimensions
    label_image = labels_spectral.reshape(H, W)

    # Create 1 by 2 plot for comparison
    fig, axes = plt.subplots(1, 2, figsize=(10, 4))

    # Plot original image
    axes[0].imshow(satellite)
    axes[0].axis('off')
    axes[0].set_title('Original Satellite Image')

    # Plot labels from  fitted spectral clustering model
    axes[1].imshow(label_image, cmap='tab10')
    axes[1].axis('off')
    axes[1].set_title(f'Spectral Clustering Labels ({clusters} Clusters)')

    plt.show()

    return labels_spectral, label_image

labels_spectral, label_image = plot_clusters(spectral_3)
```

<!-- grade_id: p1-grade-spectral, points: 7, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

Now explore other parameters. Change the number of clusters n_clusters to 4, 5, and other values you want to try; do similar for other parameters in the SpectralClustering function. Show your results. Do the results look more or less reasonable as you change those parameters? 

## Model Comparison

### Effect of the RBF bandwidth parameter (gamma)

The RBF affinity kernel is $W_{ij} = \exp(-\gamma \|x_i - x_j\|^2)$.  
A **high gamma** means only pixels with nearly identical colours are strongly connected, fragmenting the affinity graph into many small local neighbourhoods. Below we test `gamma = 200` (vs the baseline `gamma = 50`) to observe this effect.

```python
# High gamma tightens the affinity kernel — observe the effect vs baseline gamma=50
spectral_3_ga_5 = SpectralClustering(n_clusters=3, affinity='rbf', gamma=200.0, random_state=0, n_init=5)
_ = plot_clusters(spectral_3_ga_5)
```

### Effect of the number of clusters (n_clusters)

We now vary `n_clusters` from 2 to 6 (and briefly up to 25) while keeping `gamma = 50` fixed, to see how the segmentation evolves as we ask the model to resolve finer distinctions in the image.

```python
spectral_2 = SpectralClustering(n_clusters=2, affinity='rbf', gamma=50.0, random_state=0, n_init=5)
spectral_4 = SpectralClustering(n_clusters=4, affinity='rbf', gamma=50.0, random_state=0, n_init=5)
spectral_5 = SpectralClustering(n_clusters=5, affinity='rbf', gamma=50.0, random_state=0, n_init=5)
spectral_6 = SpectralClustering(n_clusters=6, affinity='rbf', gamma=50.0, random_state=0, n_init=5)

# spectral_3 (baseline, gamma=50) was already plotted above; show the remaining models
# return values assigned to "_" to avoid printing the labels + images in the output
_ = plot_clusters(spectral_2)
_ = plot_clusters(spectral_4)
_ = plot_clusters(spectral_5)
_ = plot_clusters(spectral_6)
```

```python
spectral = SpectralClustering(n_clusters=25, affinity='rbf', gamma=50.0, random_state=0, n_init=5)
_ = plot_clusters(spectral)
```

### Discussion

**Effect of `n_clusters`:**  
With `n_clusters = 2` the image is split into two broad regions (roughly light vs dark), losing meaningful distinctions (underfit). At `n_clusters = 3` (the baseline) three visually coherent zones emerge - this aligns well with the land-cover types in the image. Increasing to 4 or 5 reveals finer sub-regions (e.g. distinguishing different vegetation shades), which still look geographically plausible. At `n_clusters = 6` we see some regions begin to fragment in ways that appear arbitrary rather than reflecting genuine scene boundaries. At 25 clusters the segmentation is clearly over-partitioned to noise.

**Effect of `gamma`:**  
Raising `gamma` from 50 to 200 tightens the affinity kernel so that only very closely coloured pixels are linked. With 3 clusters this produces a spottier, less spatially coherent segmentation - regions that should belong together are separated because small colour gradients now break their connections in the graph. This confirms the expected behaviour: moderate `gamma` values that match the natural colour variation in the image give the most interpretable results.

---

# Part II - PCA and Robust PCA for Image Analysis (40 marks)

\
This part has two sections. In **Section A** we apply PCA via SVD to a set of handwritten
digit images to compute *eigendigits* and study low-rank reconstructions. In **Section B**
we implement **Robust PCA** to separate a corrupted image into a clean low-rank background
and a sparse corruption layer.

The digit images are **24×24** pixels each, stored flattened as vectors of length 576.

---

## Section A — Eigendigits via PCA (20 marks)

```python
f = open('digits.npy', 'rb')
digits = np.load(f)
f.close()
nimages = len(digits)
print('Dataset shape:', digits.shape)

def vec2img(vec):
    return vec.reshape(24, 24)
```

Randomly select **1000** images from the dataset and work with these throughout Section A.
Store them in **<font color='red'>X_sample</font>**.

1. Compute **<font color='red'>mean_digit</font>** — the mean image (shape `(576,)`) of the 1000 samples.
2. Subtract the mean to get **<font color='red'>X_centered</font>** (shape `(1000, 576)`).
3. Compute the SVD of **X_centered** using `svd(..., full_matrices=False)`. Store the results
   in **<font color='red'>U</font>**, **<font color='red'>S_vals</font>**, **<font color='red'>Vt</font>**.
   The rows of **Vt** are the **eigendigits**.

**Important:** use `np.random.seed(0)` before sampling so results are reproducible.

<!-- grade_id: p2-grade0, points: 0, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

<!-- grade_id: p2-code-pca, points: 0, grade: false, solution: true -->
```python
np.random.seed(0)

# Select 1000 images randomly and flatten them
indices = np.random.choice(nimages, 1000, replace=False)
X_sample = digits[indices].reshape(1000, -1).astype(np.float64)

# Compute mean and center
mean_digit = X_sample.mean(axis=0)
X_centered = X_sample - mean_digit

# SVD — full_matrices=False for memory efficiency
U, S_vals, Vt = svd(X_centered, full_matrices=False)

```

<!-- grade_id: p2-grade-pca, points: 5, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

assert mean_digit.shape == (576,)
assert X_centered.shape == (1000, 576)
assert_almost_equal(X_centered.mean(axis=0), np.zeros(576), decimal=10)
assert Vt.shape[1] == 576
print('--- TEST PASSED ---')
```

Display **mean_digit** and the **first 16 eigendigits** in a single row of 16 subplots.
Use `vec2img` to reshape each vector before displaying with `plt.imshow(..., cmap='gray')`.

<!-- grade_id: p2-code-show, points: 0, grade: false, solution: true -->
```python
# Create a figure with 16 subplots in a single row
fig, axes = plt.subplots(1, 16, figsize=(20, 2))

# Display the mean digit in the first subplot
axes[0].imshow(vec2img(mean_digit), cmap='gray')
axes[0].set_title('Mean Digit')
axes[0].axis('off')

# Display the first 15 eigendigits in the remaining subplots
for i in range(15):
    axes[i + 1].imshow(vec2img(Vt[i, :]), cmap='gray')
    axes[i + 1].set_title(f'Eigendigit {i + 1}')
    axes[i + 1].axis('off')

plt.tight_layout()
plt.show()

```

<!-- grade_id: p2-grade-show, points: 3, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

Reconstruct the **first image** in `X_sample` using k = 2, 5, 20, 50, 100 eigendigits.
For each k, compute the reconstruction as:
$$\hat{x} = \text{mean\_digit} + (x_{\text{centered}} \cdot V_k^\top) \cdot V_k$$
where $V_k$ contains the first k rows of **Vt**.

Display the original and the 5 reconstructions in a single row. Store the reconstruction
for k=50 in **<font color='red'>recon_50</font>**.

<!-- grade_id: p2-code-reconstruct, points: 0, grade: false, solution: true -->
```python
# Get the first centered image
x_centered = X_centered[0]

# Define the values of k for reconstruction
k_values = [2, 5, 20, 50, 100]

# Compute reconstructions for each k
reconstructions = {}
for k in k_values:
    # Extract the first k eigendigits (rows of Vt)
    V_k = Vt[:k, :]
    
    # Project the centered image onto the eigendigits
    projection = x_centered @ V_k.T
    
    # Reconstruct: mean_digit + projection * V_k
    reconstruction = mean_digit + projection @ V_k
    reconstructions[k] = reconstruction

# Store the reconstruction for k=50
recon_50 = reconstructions[50]

# Display original and reconstructions in a single row
fig, axes = plt.subplots(1, 6, figsize=(15, 2.5))

# Display original image
axes[0].imshow(vec2img(X_sample[0]), cmap='gray')
axes[0].set_title('Original')
axes[0].axis('off')

# Display reconstructions for each k
for idx, k in enumerate(k_values):
    axes[idx + 1].imshow(vec2img(reconstructions[k]), cmap='gray')
    axes[idx + 1].set_title(f'k={k}')
    axes[idx + 1].axis('off')

plt.tight_layout()
plt.show()

```

<!-- grade_id: p2-grade-reconstruct, points: 3, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

assert recon_50.shape == (576,)
print('--- TEST PASSED ---')
```

Plot the **cumulative explained variance** against the number of components (use all
singular values in `S_vals`). The explained variance ratio for component $i$ is
$s_i^2 / \sum_j s_j^2$.

Store the cumulative explained variance array in **<font color='red'>cum_var</font>**.
In the Markdown cell below state: how many eigendigits are needed to reach **85%** explained
variance?

<!-- grade_id: p2-code-variance, points: 0, grade: false, solution: true -->
```python
# Compute the explained variance ratio for each component
explained_var_ratio = (S_vals ** 2) / np.sum(S_vals ** 2)

# Compute cumulative explained variance
cum_var = np.cumsum(explained_var_ratio)

# Plot the cumulative explained variance
plt.figure(figsize=(10, 6))
plt.plot(range(1, len(cum_var) + 1), cum_var, 'b-', linewidth=2)
plt.axhline(y=0.85, color='r', linestyle='--', label='85% threshold')
plt.xlabel('Number of Components (Eigendigits)')
plt.ylabel('Cumulative Explained Variance')
plt.title('Cumulative Explained Variance vs Number of Components')
plt.grid(True, alpha=0.3)
plt.legend()
plt.xlim(0, len(cum_var))
plt.ylim(0, 1)
plt.show()

# Find the number of eigendigits needed for 85% explained variance
num_components_85 = np.argmax(cum_var >= 0.85) + 1
print(f'Number of eigendigits needed for 85% explained variance: {num_components_85}')

```

<!-- grade_id: p2-grade-variance, points: 4, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

assert len(cum_var) == len(S_vals)
assert_almost_equal(cum_var[-1], 1.0, decimal=5)
print('--- TEST PASSED ---')
```

**15 eigendigits** are needed to reach 85% explained variance.

---

## Section B — Robust PCA via IALM (20 marks)

\
We now implement **Robust PCA** to decompose a corrupted image $M = L + S$ into a
**low-rank component** $L$ and a **sparse corruption** $S$, by solving:
$$\min_{L,S}\; \|L\|_* + \lambda\|S\|_1 \quad\text{s.t.}\quad L+S=M.$$

We use the **Inexact Augmented Lagrangian Method (IALM)**:
$$\begin{aligned}
L^{(k+1)} &= \mathcal{D}_{1/\mu}(M - S^{(k)} + Y^{(k)}/\mu),\\
S^{(k+1)} &= \mathcal{T}_{\lambda/\mu}(M - L^{(k+1)} + Y^{(k)}/\mu),\\
Y^{(k+1)} &= Y^{(k)} + \mu(M - L^{(k+1)} - S^{(k+1)}),\\
\mu^{(k+1)} &= \min(\rho\,\mu^{(k)},\;\mu_{\max}),
\end{aligned}$$
where $\mathcal{D}_\tau$ soft-thresholds singular values and
$\mathcal{T}_\epsilon(x) = \text{sign}(x)\max(|x|-\epsilon,0)$ is element-wise soft-thresholding.

Stop when $\|M - L^{(k)} - S^{(k)}\|_F / \|M\|_F \le \text{tolerance}$.

The penalty parameter is initialized following Lin et al. (2010):
$$\mu_0 = 1.25 \;/\; \|M\|_2,$$
where $\|M\|_2$ is the spectral norm (largest singular value of $M$).
This ties the initial SVT threshold directly to the scale of the data,
ensuring the low-rank and sparse components are separated correctly from
the very first iteration.

**Step 1: Implement the two thresholding operators.**

- **<font color='red'>singular_value_threshold(A, tau)</font>** — compute SVD of A,
  soft-threshold singular values by tau, return reconstructed matrix.
- **<font color='red'>element_threshold(A, eps)</font>** — element-wise soft-threshold:
  $\text{sign}(A)\max(|A|-\varepsilon,\,0)$.

<!-- grade_id: p2-code-thresh, points: 0, grade: false, solution: true -->
```python
def element_threshold(A, eps):
    A = np.array(A, dtype=float)
    return np.sign(A) * np.maximum(np.abs(A) - eps, 0.0)


def singular_value_threshold(A, tau):
    A = np.array(A, dtype=float)
    U, s, Vt = svd(A, full_matrices=False)
    s_thr = np.maximum(s - tau, 0.0)
    return (U * s_thr) @ Vt
```

<!-- grade_id: p2-grade-thresh, points: 5, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

# element_threshold
assert_almost_equal(element_threshold(np.array([[0.3,-0.2],[1.0,-1.5]]),0.5),
                    np.array([[0.0,0.0],[0.5,-1.0]]))
# singular_value_threshold: rank should not increase
B = np.outer([1.0,2.0,3.0],[1.0,0.5])  # rank-1
D = singular_value_threshold(B, 0.1)
assert D.shape == B.shape
# large tau => zero matrix
assert_almost_equal(singular_value_threshold(np.eye(3)*0.1, 1.0), np.zeros((3,3)))
print('--- TEST PASSED ---')
```

**Step 2: Implement the IALM Robust PCA solver.**

Write **<font color='red'>robust_pca(M, lam, mu, rho, mu_max, tolerance, max_iter, iter_print)</font>**
with defaults:
- `lam = None` → $1/\sqrt{\max(D,N)}$
- `mu = None`  → $1.25 \;/\; \|M\|_2$ &nbsp; (use `svd(M, compute_uv=False)[0]` for $\|M\|_2$)
- `mu_max = None` → $10^7 \times \mu_0$
- `rho = 1.5`, `tolerance = 1e-5`, `max_iter = 500`, `iter_print = 50`

Return **(L, S)**.

<!-- grade_id: p2-code-rpca, points: 0, grade: false, solution: true -->
```python
def robust_pca(M, lam=None, mu=None, rho=1.5, mu_max=None,
               tolerance=1e-5, max_iter=500, iter_print=50):
    M = np.array(M, dtype=float)
    D, N = M.shape

    if lam is None:
        lam = 1.0 / np.sqrt(max(D, N))
    if mu is None:
        mu = 1.25 / svd(M, compute_uv=False)[0]
    if mu_max is None:
        mu_max = 1e7 * mu

    L = np.zeros_like(M)
    S = np.zeros_like(M)
    Y = M / max(norm(M, 'fro'), np.finfo(float).eps)
    norm_M = norm(M, 'fro')

    for k in range(max_iter):
        L = singular_value_threshold(M - S + Y / mu, 1.0 / mu)
        S = element_threshold(M - L + Y / mu, lam / mu)
        residual = M - L - S
        residual_norm = norm(residual, 'fro')
        if residual_norm / norm_M <= tolerance:
            break
        Y = Y + mu * residual
        mu = min(rho * mu, mu_max)

    return L, S
```

The cell below tests your robust_pca implementation on a small synthetic matrix of size 20 × 25. It constructs a rank-2 ground-truth low-rank component L and a sparse corruption S with 10% non-zero entries set to ±5, then checks that your decomposition recovers L+S=M with relative error below 1%. Complete the code. 

<!-- grade_id: p2-grade-rpca, points: 5, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

np.random.seed(123)
D_r, N_r = 20, 25
U_r = np.random.randn(D_r, 2); V_r = np.random.randn(2, N_r)
L_true = U_r @ V_r
S_true = np.zeros((D_r, N_r))
idx_c = np.random.choice(D_r*N_r, size=int(0.10*D_r*N_r), replace=False)
S_true.flat[idx_c] = np.random.choice([-5.0, 5.0], size=len(idx_c))
M_t = L_true + S_true
L_r, S_r = robust_pca(M_t)

rel_err = norm(M_t - L_r - S_r, 'fro') / norm(M_t, 'fro')
assert rel_err < 0.01, f'Error too large: {rel_err:.4f}'
print(f'Relative error: {rel_err:.6f}')
print('--- TEST PASSED ---')
```

**Step 3: Denoise a corrupted image.**

Load the astronaut image, convert to grayscale, and corrupt **25%** of pixels by setting
them to 0 or 1 at random. Store the corrupted image in
**<font color='red'>corrupted_gray</font>**.

Apply `robust_pca` (tolerance=1e-4, max_iter=150) and clip the recovered L to [0,1].
Store in **<font color='red'>L_recovered</font>**.

Plot original, corrupted, and recovered side by side.

<!-- grade_id: p2-code-apply, points: 0, grade: false, solution: true -->
```python
# Load and plot the original colorful image 
astro_color = img_as_float(data.astronaut())
# Plot the colorful
plt.imshow(astro_color)
plt.axis('off')
plt.title('Colorful Original')
plt.show()


```

```python
# Convert to grayscale and plot the image
astro_gray  = rgb2gray(astro_color)   # shape (512, 512)

# Plot the grayscale
plt.imshow(astro_gray, cmap='gray')
plt.axis('off')
plt.title('Grayscale Original')
plt.show()
```

```python
# Corrupt 25% of pixels
np.random.seed(123)
corrupted_gray = astro_gray.copy()
num_pixels = corrupted_gray.size
idx = np.random.choice(num_pixels, size=int(0.25 * num_pixels), replace=False)
corrupted_gray.flat[idx] = np.random.choice([0.0, 1.0], size=len(idx))
```

```python
# Apply Robust PCA with the specified parameters
L_astro, S_astro = robust_pca(corrupted_gray, tolerance=1e-4, max_iter=150)

# Clip recovered low-rank component to valid pixel range [0, 1]
L_recovered = np.clip(L_astro, 0.0, 1.0)

# Plot original, corrupted, and recovered side by side
fig, axes = plt.subplots(1, 3, figsize=(15, 5))

axes[0].imshow(astro_gray, cmap='gray')
axes[0].axis('off')
axes[0].set_title('Original Grayscale')

axes[1].imshow(corrupted_gray, cmap='gray')
axes[1].axis('off')
axes[1].set_title('Corrupted (25% pixels)')

axes[2].imshow(L_recovered, cmap='gray')
axes[2].axis('off')
axes[2].set_title('Recovered (Low-Rank L)')

plt.tight_layout()
plt.show()
```

<!-- grade_id: p2-grade-apply, points: 3, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

In the Markdown cell below answer:
1. What does the low-rank component $L$ represent, and what does sparse component $S$ represent?
2. How does Robust PCA differ from standard PCA when outliers are present?
3. What happens qualitatively if $\lambda$ is set much larger or much smaller?

Answers:
1. The low-rank component $L$ is the recovered structured image. It is the result of projecting the corrupted image onto a low-dimensional structure/manifold where the "true" image lives. Component S represents the left over anomaly $S = X - L$. These are the pixels that do not fit the coherent structure of the image (localised anomolies, corruptions).

2. Standard PCA is extremely brittle in the presence of large outliers. This is because standard PCA looks for the best rank-$k$ estimate of the low-rank $L$ via:
    $$
    \begin{aligned}
    \text{minimize} \quad & \|X-L\|_F \\
    \text{subject to} \quad & \operatorname{rank}(L) \le k
    \end{aligned}
    $$

    Where it's clear that large anomalies dominate.

    Robust PCA considers a different optimisation problem:
    $$
    \begin{aligned}
    \min_{L,S} \quad & \|L\|_* + \lambda \|S\|_1 \\
    \text{subject to} \quad & X = L + S
    \end{aligned}
    $$

    Separating the data into a low-rank and sparse components. Which performs much better with outliers since a single low-rank model isn't forced to explain large anomalous data. 

3. Setting $\lambda$ larger leads to more noise remaining in low-rank $L$, hence resembling less the original image. The sparse image $S$ also becomes even sparser, keeping only the strongest anomalies.

    Setting $\lambda$ **smaller** leads to more of the true underlying image being mistaken for noise, and so moved into $S$.

---

# Part III - Autoencoder for Dimensionality Reduction (30 marks)

\
We implement a **shallow autoencoder** from scratch using only NumPy and apply it to
compress and reconstruct handwritten digit images.

The network has the structure:
$$z = \sigma(W_1 x + b_1), \qquad \hat{x} = W_2 z + b_2,$$
where $\sigma$ is the **sigmoid** activation applied element-wise and the output layer is
**linear**. Parameters: $W_1\in\mathbb{R}^{p\times d}$, $b_1\in\mathbb{R}^p$,
$W_2\in\mathbb{R}^{d\times p}$, $b_2\in\mathbb{R}^d$.

Loss: $\mathcal{L}=\frac{1}{n}\|{\hat{X}}-X\|_F^2$. Optimized via **mini-batch gradient descent**.

---

```python
f = open('digits.npy','rb')
digits_all = np.load(f); f.close()
N_train = 500
X_train = digits_all[:N_train].astype(float).reshape(N_train, -1) 
print('Training data shape:', X_train.shape)
```

**Step 1: Sigmoid and its derivative.**

- **<font color='red'>sigmoid(z)</font>** — element-wise $1/(1+e^{-z})$.
- **<font color='red'>sigmoid_deriv(z)</font>** — element-wise $\sigma(z)(1-\sigma(z))$.

```python
def sigmoid(z):
    return 1.0 / (1.0 + np.exp(-z))

def sigmoid_deriv(z):
    sig = sigmoid(z)
    return sig * (1.0 - sig)
```

```python
### DO NOT REMOVE/EDIT THIS CELL ###

z_t = np.array([-2.0, 0.0, 2.0])
assert_almost_equal(sigmoid(z_t), [0.11920292, 0.5, 0.88079708], decimal=6)
assert_almost_equal(sigmoid_deriv(z_t), sigmoid(z_t)*(1-sigmoid(z_t)), decimal=10)
assert_almost_equal(sigmoid_deriv(np.array([0.0])), [0.25], decimal=8)
print('--- TEST PASSED ---')
```

\
**Step 2: Forward pass.**

**<font color='red'>forward(X, W1, b1, W2, b2)</font>** — takes input batch `(n, d)` and returns
**(z, x_hat)**:
- $z = \sigma(X W_1^\top + b_1)$, shape `(n, p)`.
- $\hat{x} = z W_2^\top + b_2$, shape `(n, d)`.

```python
def forward(X, W1, b1, W2, b2):
    # (n, p) matrix of all encoded samples.
    z = sigmoid(X @ W1.T + b1)

    # (n, d) matrix of all reconstructed samples.
    x_hat = z @ W2.T + b2
    
    return z, x_hat
```

```python
### DO NOT REMOVE/EDIT THIS CELL ###

np.random.seed(1)
d_t,p_t,n_t = 576,32,10
W1_t=np.random.randn(p_t,d_t)*0.01; b1_t=np.zeros(p_t)
W2_t=np.random.randn(d_t,p_t)*0.01; b2_t=np.zeros(d_t)
X_t=np.random.rand(n_t,d_t)
z_t,xh_t=forward(X_t,W1_t,b1_t,W2_t,b2_t)
assert z_t.shape==(n_t,p_t) and xh_t.shape==(n_t,d_t)
assert z_t.min()>0 and z_t.max()<1
print('--- TEST PASSED ---')
```

\
**Step 3: Backward pass.**

**<font color='red'>backward(X, z, x_hat, W1, b1, W2, b2)</font>** — returns gradients
**(dW1, db1, dW2, db2)** of $\mathcal{L}=\frac{1}{n}\|\hat{X}-X\|_F^2$.

*Hint: start from the output layer error $\delta_2 = \frac{2}{n}(\hat{X}-X)$, then
backpropagate through $W_2$, through the sigmoid, and through $W_1$.*

```python
def backward(X, z, x_hat, W1, b1, W2, b2):
    n = X.shape[0]
    
    # Output layer error: dL/d(x_hat)
    delta_2 = (2.0 / n) * (x_hat - X)  # shape (n, d)
    
    # Gradients for W2 and b2
    # x_hat = z @ W2.T  =>  dL/dW2 = (dL/d(x_hat)).T @ z
    dW2 = delta_2.T @ z              # shape (d, p)
    db2 = np.sum(delta_2, axis=0)   # shape (d,)
    
    # Backpropagate to hidden layer
    # x_hat = z @ W2.T  =>  dL/dz = dL/d(x_hat) @ W2  (not W2.T)
    dz = delta_2 @ W2               # shape (n, p)
    delta_1 = dz * z * (1 - z)     # sigmoid derivative: sigma*(1-sigma)
    
    # Gradients for W1 and b1
    # z = sigmoid(X @ W1.T + b1)  =>  dL/dW1 = delta_1.T @ X
    dW1 = delta_1.T @ X             # shape (p, d)
    db1 = np.sum(delta_1, axis=0)  # shape (p,)
    
    return dW1, db1, dW2, db2
```

**Step 4: Training loop.**

**<font color='red'>train_autoencoder(X, p, lr, n_epochs, batch_size, random_state)</font>**:
- `lr=0.01`, `n_epochs=50`, `batch_size=128`, `random_state=0`
- Initialize weights: Gaussian std=0.01, biases=0.
- Return **(W1, b1, W2, b2, loss_history)**.

After training: plot the loss curve and display **5 original vs 5 reconstructed** digit images
side by side.

```python
def train_autoencoder(X, p, lr=0.01, n_epochs=50, batch_size=128, random_state=0):
    np.random.seed(random_state)
    
    n, d = X.shape
    
    # Initialise weights with Gaussian std=0.01, biases to zero
    W1 = np.random.randn(p, d) * 0.01
    b1 = np.zeros(p)
    W2 = np.random.randn(d, p) * 0.01
    b2 = np.zeros(d)
    
    loss_history = []
    
    for epoch in range(n_epochs):
        # Shuffle data each epoch
        indices = np.random.permutation(n)
        X_shuffled = X[indices]
        
        epoch_loss = 0.0
        
        # Mini-batch gradient descent
        for i in range(0, n, batch_size):
            X_batch = X_shuffled[i:i + batch_size]
            
            z, x_hat = forward(X_batch, W1, b1, W2, b2)
            
            epoch_loss += np.mean((x_hat - X_batch) ** 2) * X_batch.shape[0]
            
            dW1, db1, dW2, db2 = backward(X_batch, z, x_hat, W1, b1, W2, b2)
            
            W1 -= lr * dW1
            b1 -= lr * db1
            W2 -= lr * dW2
            b2 -= lr * db2
        
        avg_loss = epoch_loss / n
        loss_history.append(avg_loss)
        
        if (epoch + 1) % 10 == 0:
            print(f'Epoch {epoch+1}/{n_epochs}, Loss: {avg_loss:.6f}')
    
    # Plot loss curve
    plt.figure(figsize=(8, 5))
    plt.plot(loss_history, linewidth=2)
    plt.xlabel('Epoch')
    plt.ylabel('Loss')
    plt.title('Autoencoder Training Loss')
    plt.grid(True, alpha=0.3)
    plt.tight_layout()
    plt.show()
    
    # Display 5 original vs 5 reconstructed digits
    _, x_hat_test = forward(X[:5], W1, b1, W2, b2)
    fig, axes = plt.subplots(2, 5, figsize=(12, 5))
    for i in range(5):
        axes[0, i].imshow(X[i].reshape(24, 24), cmap='gray')
        axes[0, i].set_title(f'Original {i}')
        axes[0, i].axis('off')
        axes[1, i].imshow(x_hat_test[i].reshape(24, 24), cmap='gray')
        axes[1, i].set_title(f'Reconstructed {i}')
        axes[1, i].axis('off')
    plt.tight_layout()
    plt.show()
    
    return W1, b1, W2, b2, loss_history
```

```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

**Step 5: Compare with PCA.**

Using SVD on `X_train`, compute PCA reconstructions with **p = 32** components
(same bottleneck as the autoencoder). Display original, PCA reconstruction, and autoencoder
reconstruction for the same 5 images in three rows.

In the Markdown cell below answer:
1. Which method gives better visual reconstructions? Does this match your expectation?
2. What objective function difference explains the quality gap?
3. A **linear autoencoder** (identity activation instead of sigmoid) is mathematically
   equivalent to which method, and why?

```python
# Train the autoencoder with p=32
p = 32
W1, b1, W2, b2, loss_history = train_autoencoder(X_train, p, lr=0.01, n_epochs=200, batch_size=128, random_state=0)

# Compute PCA reconstructions with p=32 components using the admitted svd import
X_mean = np.mean(X_train, axis=0)
X_centered = X_train - X_mean
U_pca, S_pca, Vt_pca = svd(X_centered, full_matrices=False)

# Project onto first p principal components and reconstruct
projection_pca = X_centered @ Vt_pca[:p].T  # shape (n, p)
X_recon_pca = projection_pca @ Vt_pca[:p] + X_mean

# Get autoencoder reconstructions
z_ae, x_hat_ae = forward(X_train, W1, b1, W2, b2)

# Display original, PCA, and autoencoder reconstructions for 5 images
fig, axes = plt.subplots(3, 5, figsize=(12, 8))

for i in range(5):
    axes[0, i].imshow(X_train[i].reshape(24, 24), cmap='gray')
    axes[0, i].set_title(f'Original {i}', fontsize=10)
    axes[0, i].axis('off')

    axes[1, i].imshow(X_recon_pca[i].reshape(24, 24), cmap='gray')
    axes[1, i].set_title(f'PCA {i}', fontsize=10)
    axes[1, i].axis('off')

    axes[2, i].imshow(x_hat_ae[i].reshape(24, 24), cmap='gray')
    axes[2, i].set_title(f'Autoencoder {i}', fontsize=10)
    axes[2, i].axis('off')

plt.tight_layout()
plt.show()

# Compute and report reconstruction errors
pca_error = np.mean((X_recon_pca - X_train) ** 2)
ae_error = np.mean((x_hat_ae - X_train) ** 2)

print(f'PCA Reconstruction MSE: {pca_error:.6f}')
print(f'Autoencoder Reconstruction MSE: {ae_error:.6f}')
```

```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

## Comparison of Autoencoder with PCA

**1. Which method gives better visual reconstructions? Does this match your expectation?**

PCA achieves a lower reconstruction MSE (see printed values above), indicating better performance for this dataset. For digit images with p=32 components, PCA provides better reconstructions as expected: handwritten digits live close to a low-dimensional linear subspace, so the globally optimal linear projection found by PCA is hard to beat. The autoencoder still produces recognisable reconstructions but converges to a local minimum with higher reconstruction error.

**2. What objective function difference explains the quality gap?**

PCA minimizes the reconstruction error over all rank-32 linear approximations of the centered data, finding the globally optimal linear subspace via SVD. The autoencoder minimizes the same squared-error loss but through gradient descent on a nonlinear sigmoid encoder, which can get stuck in local minima - and with only 500 training samples and 200 epochs, it hasn't really found a good one. While the nonlinearity in principle provides more representational capacity, it also makes the optimisation harder - PCA's closed-form solution is unreachable by the constrained nonlinear network.

**3. A linear autoencoder (identity activation instead of sigmoid) is mathematically equivalent to which method, and why?**

A linear autoencoder with squared-error loss is mathematically equivalent to PCA. Because the encoder and decoder are both linear, their composition is also linear and can only represent a linear subspace of dimension at most $p$. Minimising the reconstruction loss over all such networks is therefore the same problem as finding the best rank-$p$ approximation of the centered data - which is exactly what PCA solves via the SVD.