# MTH793P - Coursework 1

This is a template notebook for the computational exercises of Coursework #1 of the module MTH793P, Advanced Machine Learning. 

```python
import numpy as np
import matplotlib.pyplot as plt
from scipy import sparse as sp
from scipy.sparse import linalg
from numpy.testing import assert_array_equal
```

Create two lists *nodes* and *edges* and one NumPy array *weights*. The list *nodes* should contain all names of the nodes in the graph of Coursework 1, Question 1, in alphabetical order. The list edges should include lists that contain the indices of the nodes that are connected by the individual edge. For example, the first edge connects node 'Batman' and node 'Jessica Jones', so the list for this edge should be [0, 3], as 'Batman' is the first entry of the list *nodes* and 'Jessica Jones' is the fourth entry.  The Numpy array *weights* should contain the individual weights assigned to each edge.

<!-- grade_id: cell-1b50dfad57619207, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
```

Display your lists and array with the following cell.

<!-- grade_id: cell-7aaa793b98db93ce, points: 3, grade: true, solution: false -->
```python
print('We consider a graph with nodes/vertices {n}, edges {e} and weights {w}.'.format( \
                            n = nodes, e = edges, w = weights))
```

Write a function **construct_incidence_matrix** that takes the lists *nodes*, *edges* and the NumPy array *weights* as arguments. The function should return the incidence matrix *incidence_matrix* that corresponds to the weighted graph. The construction of the incidence matrix is described in the lecture notes.

<!-- grade_id: cell-87385a31ebe6ae36, points: 0, grade: false, solution: true -->
```python
def construct_incidence_matrix(nodes, edges, weights):
    ### YOUR CODE HERE
```

You can test your function with the following cell.

<!-- grade_id: cell-81c3933d937a234f, points: 4, grade: true, solution: false -->
```python
incidence_matrix = construct_incidence_matrix(nodes, edges, weights)    
print('The incidence matrix of our graph is \n {i}.'.format(i = incidence_matrix.astype(int)))

### TEST YOUR CODE HERE
test_nodes = ['Batman', 'Catwoman', 'Spiderman']
test_edges = [[0, 1], [1, 2]]
test_weights = np.array([81, 1])
test_incidence_matrix = construct_incidence_matrix(test_nodes, test_edges, test_weights)


assert_array_equal(test_incidence_matrix.astype(int), np.array([[-9, 9, 0],[0, -1, 1]]))
assert_array_equal(incidence_matrix.astype(int), np.array([[-2, 0, 0, 2, 0, 0],[-9, 9, 0, 0, 0, 0],[-4, 0, 4, 0, 0, 0],[-8, 0, 0, 0, 0, 8],[0, 0, -8, 8, 0, 0],[0, 0, 0, -6, 0, 6],[0, 0, 0, -7, 7, 0],[0, -7, 0, 0, 0, 7],[0, -1, 0, 0, 1, 0],[0, 0, -7, 0, 7, 0],[0, 0, 0, 0, -1, 1]]))

```

Next, compute the corresponding graph-Laplacian for the incidence matrix *incidence_matrix* from the previous exercise and store it in a variable named *graph_laplacian*. Follow the definition from the lecture notes.

<!-- grade_id: cell-bbe113db448d0f54, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
```

You can test your function with the following cell.

<!-- grade_id: cell-e14e6d8159b7b419, points: 3, grade: true, solution: false -->
```python
### TEST YOUR CODE HERE
assert_array_equal(graph_laplacian.astype(int), incidence_matrix.T @ incidence_matrix)
test_graph_laplacian = test_incidence_matrix.T @ test_incidence_matrix
assert_array_equal(test_graph_laplacian.astype(int), np.array([[81, -81, 0],[-81, 82, -1],[0, -1, 1]]))

print('The graph Laplacian of our graph is \n {g}.'.format(g = graph_laplacian.astype(int)))
```

We want to use our graph to determine whether a node in the graph belongs to the class "Marvel" or the class "DC". Suppose we are in a semi-supervised setting, where the node "Deadpool" is already labelled $v_{\text{Deadpool}} = 0$ (class "Marvel") and the node "Catwoman" is labelled as $v_{\text{Catwoman}} = 1$ (class "DC"). Here $v$ is the mathematical notation of the label vector. We follow the instructions in the lecture notes and formulate this as a linear system. We can either define appropriate projection matrices or create sub-matrices from the graph-Laplacian by choosing the correct indices. How you set up the linear system is up to you. Store your linear system in a variable named *linear_system* and the right-hand-side in a variable *right_hand_side*.

<!-- grade_id: cell-20979238f798da95, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
```

Solve your linear system and store your labels in an array named *remaining_labels*. Create also an array *all_labels* that contains all labels, as well as a boolean array *thresholded_labels* of the same size as *all_labels* with True or False values. An entry should be true if the corresponding entry in *all_labels* is larger than 0.5 and false otherwise.

<!-- grade_id: cell-ae555d965960f117, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
```

Check your results with the following cell.

```python
print('The computed labels are {a}. Setting all values above 0.5 to one and'.format( \
        a = F.T), 'the remaining ones to zero yields {t}.'.format(t = \
        Labels.astype(int).T))
```

We conclude this coursework by computing the second eigenvector, i.e. the eigenvector that corresponds to the second smallest eigenvalue, of the graph-Laplacian *graph_laplacian*. Explore how to compute eigenvalues and eigenvectors, in particular of sparse matrices if you have used them.

<!-- grade_id: cell-67ec5a28d4bc7ddc, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
```

Display the second smallest eigenvalue and the corresponding eigenvector. What do you observe?

<!-- grade_id: cell-fd9c7426dfff76d8, points: 0, grade: true, solution: false -->
```python
print('The second smallest eigenvalue of the graph Laplacian is {eval}.'.format(eval = \
        eigenvalues[1].real), 'The corresponding eigenvector is {evec}.T.'.format(evec = \
        eigenvectors[:, 1].T.real))

print('If we set all entries of the second eigenvector that are larger than zero to', \
      'one and the remaining ones to zero, we obtain {t}.T.'.format(t = \
        (eigenvectors[:, 1] > 0).astype(int)), 'This yields the same classification' \
      ' as before, but without having to classify any labels in advance.')
```