# MTH793P: Advanced Machine Learning, Semester B, 2025/2026
## Final Examination

---

**Full Name:**

<!-- grade_id: hdr-fname-a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

**Student ID:**

<!-- grade_id: hdr-sid-a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

---

### General instructions

First, enter your **Full Name** and **Student ID** in the fields above.

**Read the following information carefully before turning to the actual exercises.**

This final project assignment is scored on a scale of 100 marks and will contribute **60% of your final mark** for this module.

This project is due by **Friday, 22 May 2026, 17:00 GMT**.

Submissions received after this deadline will be treated in accordance with the College Regulations for late submissions: Up to seven days after the deadline, late submissions will incur a penalty of 5 marks per 24 hours. Any submission received more than 168 hours (7 days) after the deadline will receive 0 marks.

The lecturer will be available to answer reasonable questions about this assessment until **Friday, 15 May 2026, 17:00 GMT**. You can contact the lecturer by email ([hong.qi@qmul.ac.uk](mailto:hong.qi@qmul.ac.uk)).

You must use this Jupyter Notebook to answer all exercises. You are not allowed to remove any cells from this notebook. You can add Markdown or code cells as needed.

You must submit this Jupyter Notebook through QMPlus. Files sent by other means, including email, will not be considered.

You cannot submit any additional files apart from this Jupyter Notebook. The ZIP file you downloaded contains various auxiliary files that you will need for the exercises. The original versions of these files will be merged into your submission automatically.


For **Markdown** exercises, **replace all occurrences of** `YOUR ANSWER HERE`  **by valid Markdown** to answer the question.

For **coding** exercises,
**replace all occurrences of**
```python
### YOUR CODE HERE
raise NotImplementedError()
```
**by valid Python code** to solve the problem.
In your solution, you can **only use** built-in Python functions and **functions from modules were imported for you in this template**.
Importing other modules is not allowed and will result in withdrawal of the corresponding marks.
If your code produces an error and you cannot fix it, write a comment to show that you are aware of the problem and indicate possible causes if you can.
You can get **partial credit**, so even if you cannot solve an exercise completely, try to answer as much as you can.


Finally, before you turn this project in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart Kernel) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All (Jupyter Notebook Version 6) or Run$\rightarrow$Run All Cells (Jupyter Notebook Version 7)). Please **do not change the file name** when you upload your Notebook to QMPlus.

<div class="alert alert-block alert-warning">
    Please answer all the questions below by writing your own Python code. Using or copy-pasting pre-existing code from any source is not permitted, and will be considered as <b>plagiarism</b> and addressed according to College Regulations. You are not allowed to collaborate with other students or ask for help from any other source, including, but not limited to, online forums or ChatGPT (or any other AI agent). By submitting your Notebook via QMPlus, you confirm that you have followed these rules and that the submitted code is your own work. Upon violation, you will fail this test with zero marks.</div>

---

### Coding style

#### Comments

Comments are an essential part of programming.
Use appropriate comments to explain the reasoning behind your code.

#### Variables & Functions

Use the variable and function names specified (**<font color=red>in red</font>**) in each exercise when writing your code. If you use other names, the testing of your code might fail, resulting in a reduced mark. For any additional variables/functions you define, choose meaningful, concise names.


It is perfectly fine if you define extra auxiliary functions to carry out a task. This is to say, even if the exercise only asks you to write a function `example_function`, you are free to add other functions like `example_function_aux`, `some_other_function`, ... and use them in your implementation of `example_function`.
Choose meaningful names for all auxiliary functions, too.

---

### Assessment criteria

The first thing we will do to evaluate your submission is to clear all outputs and run the entire notebook from the beginning to the end. It is your responsibility to ensure that the code produces **no errors**. We will not try to fix any errors when assessing your work as you have several weeks to develop this project. Instead, you will receive **zero marks for any exercises where the code cell raises an error**. If you run into an error and do not know how to fix it, comment out the critical line of code and any subsequent ones that depend on it. Add a remark indicating that you are aware of the problem and potential reasons.

If you do not know how to implement a certain operation in Python,
add a comment and explain in plain English (and in sufficient detail) what you would like to do.
You may receive partial credit for expedient ideas formulated in this way.

**10 marks (10% of the project mark)** will be awarded for **coding style**. This includes the use of appropriate variable names, suitable comments to explain complex sequences of operations, and general organization.
This is to say, even if your code is fully functional and solves the exercise, you may lose marks if the reasoning is not explained properly, or the code is not well-organized.

<div class="alert alert-block alert-danger">

**Running time:** You will be asked to code and use several machine learning methods that we studied in the lectures. If not implemented correctly, or used with the wrong parameters, these can take a long time to run. When you submit your notebook, you should make sure that:
* No single box in this notebook takes **longer than 30 seconds** to run on your personal computer.
* The entire notebook runs from beginning to end in **less than 3 minutes** on your personal computer.

You should make sure that your submitted notebook follows these limitations. Failing to do so will amount to a penalty of **20 marks**. 

---

### Admitted imports

You may only use built-in Python and the packages imported in the cell below. **Do not import additional modules.**

```python
import numpy as np
from numpy.linalg import svd, norm

import matplotlib.pyplot as plt

from skimage import data
from skimage import img_as_float
from skimage.color import rgb2gray

from sklearn.cluster import KMeans, SpectralClustering
from sklearn.preprocessing import StandardScaler

from numpy.testing import assert_equal, assert_array_equal
from numpy.testing import assert_almost_equal, assert_array_almost_equal
```

---

# Part I - Satellite Image Segmentation Using Spectral Clustering (20 marks)

In this part we use **Spectral Clustering** to segment a satellite image of land cover into
distinct regions — forest, water, urban area, and farmland — based on pixel color similarity.

The file **`satellite.npy`** contains a small vague satellite image of shape **(H, W, 3)**, with
RGB values normalized to $[0, 1]$. Each pixel is treated as a data point with a
3-dimensional color feature vector.

For this part you may use **SpectralClustering** from `sklearn.cluster` as
well as `numpy` for the manual steps.

---

The cell below loads the image into **`satellite`** and displays it.

```python
satellite = np.load('satellite_50x70.npy')
H, W = satellite.shape[:2]
print('Image shape:', satellite.shape)
plt.imshow(satellite); plt.axis('off'); plt.title('Satellite Image'); plt.show()
```

**Step 1: Prepare the pixel feature matrix.**

Reshape the satellite image into **<font color='red'>pixels</font>** — an array of shape
`(H*W, 3)` where each row is the RGB vector of one pixel. Store the image dimensions in
**<font color='red'>H</font>** and **<font color='red'>W</font>**.

<!-- grade_id: p1-code-reshape, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
raise NotImplementedError()

# Hint: use .reshape(-1, 3) to flatten the (H, W, 3) image into a (H*W, 3) pixel matrix
# and print the resulting shape
```

<!-- grade_id: p1-grade-reshape, points: 3, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

assert pixels.ndim == 2
assert pixels.shape[1] == 3
assert pixels.shape[0] == H * W
assert_almost_equal(pixels.min(), 0.0, decimal=3)
assert pixels.max() <= 1.0 + 1e-6
print('--- TEST PASSED ---')
```

**Step 2: Apply Spectral Clustering.**

Apply **SpectralClustering** to **`pixels`** with:
- `n_clusters = 3`
- `affinity = 'rbf'`
- `gamma = 50.0`
- `random_state = 0`
- `n_init = 5`

Store the integer labels (0–3) in **<font color='red'>labels_spectral</font>**, then reshape
to **(H, W)** and store in **<font color='red'>label_image</font>**.

Display `label_image` with `cmap='tab10'` alongside the original image.

<!-- grade_id: p1-code-spectral, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
raise NotImplementedError()


```

<!-- grade_id: p1-grade-spectral, points: 7, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

Now explore other parameters. Change the number of clusters n_clusters to 4, 5, and other values you want to try; do similar for other parameters in the SpectralClustering function. Show your results. Do the results look more or less reasonable as you change those parameters? 

YOUR ANSWER HERE

---

# Part II - PCA and Robust PCA for Image Analysis (40 marks)

\
This part has two sections. In **Section A** we apply PCA via SVD to a set of handwritten
digit images to compute *eigendigits* and study low-rank reconstructions. In **Section B**
we implement **Robust PCA** to separate a corrupted image into a clean low-rank background
and a sparse corruption layer.

The digit images are **24×24** pixels each, stored flattened as vectors of length 576.

---

## Section A — Eigendigits via PCA (20 marks)

```python
f = open('digits.npy', 'rb')
digits = np.load(f)
f.close()
nimages = len(digits)
print('Dataset shape:', digits.shape)

def vec2img(vec):
    return vec.reshape(24, 24)
```

Randomly select **1000** images from the dataset and work with these throughout Section A.
Store them in **<font color='red'>X_sample</font>**.

1. Compute **<font color='red'>mean_digit</font>** — the mean image (shape `(576,)`) of the 1000 samples.
2. Subtract the mean to get **<font color='red'>X_centered</font>** (shape `(1000, 576)`).
3. Compute the SVD of **X_centered** using `svd(..., full_matrices=False)`. Store the results
   in **<font color='red'>U</font>**, **<font color='red'>S_vals</font>**, **<font color='red'>Vt</font>**.
   The rows of **Vt** are the **eigendigits**.

**Important:** use `np.random.seed(0)` before sampling so results are reproducible.

<!-- grade_id: p2-grade0, points: 0, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

<!-- grade_id: p2-code-pca, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
raise NotImplementedError()

# Hint: select 1000 images with np.random.seed(0) and np.random.choice;
#        flatten each image with .reshape(1000, -1) to get shape (1000, 576)
### YOUR CODE HERE

# Compute mean and center
### YOUR CODE HERE

# SVD — full_matrices=False is important for memory
### YOUR CODE HERE

```

<!-- grade_id: p2-grade-pca, points: 5, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

assert mean_digit.shape == (576,)
assert X_centered.shape == (1000, 576)
assert_almost_equal(X_centered.mean(axis=0), np.zeros(576), decimal=10)
assert Vt.shape[1] == 576
print('--- TEST PASSED ---')
```

Display **mean_digit** and the **first 16 eigendigits** in a single row of 16 subplots.
Use `vec2img` to reshape each vector before displaying with `plt.imshow(..., cmap='gray')`.

<!-- grade_id: p2-code-show, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: p2-grade-show, points: 3, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

Reconstruct the **first image** in `X_sample` using k = 2, 5, 20, 50, 100 eigendigits.
For each k, compute the reconstruction as:
$$\hat{x} = \text{mean\_digit} + (x_{\text{centered}} \cdot V_k^\top) \cdot V_k$$
where $V_k$ contains the first k rows of **Vt**.

Display the original and the 5 reconstructions in a single row. Store the reconstruction
for k=50 in **<font color='red'>recon_50</font>**.

<!-- grade_id: p2-code-reconstruct, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
raise NotImplementedError()


```

<!-- grade_id: p2-grade-reconstruct, points: 3, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

assert recon_50.shape == (576,)
print('--- TEST PASSED ---')
```

Plot the **cumulative explained variance** against the number of components (use all
singular values in `S_vals`). The explained variance ratio for component $i$ is
$s_i^2 / \sum_j s_j^2$.

Store the cumulative explained variance array in **<font color='red'>cum_var</font>**.
In the Markdown cell below state: how many eigendigits are needed to reach **85%** explained
variance?

<!-- grade_id: p2-code-variance, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: p2-grade-variance, points: 4, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

assert len(cum_var) == len(S_vals)
assert_almost_equal(cum_var[-1], 1.0, decimal=5)
print('--- TEST PASSED ---')
```

YOUR ANSWER HERE

---

## Section B — Robust PCA via IALM (20 marks)

\
We now implement **Robust PCA** to decompose a corrupted image $M = L + S$ into a
**low-rank component** $L$ and a **sparse corruption** $S$, by solving:
$$\min_{L,S}\; \|L\|_* + \lambda\|S\|_1 \quad\text{s.t.}\quad L+S=M.$$

We use the **Inexact Augmented Lagrangian Method (IALM)**:
$$\begin{aligned}
L^{(k+1)} &= \mathcal{D}_{1/\mu}(M - S^{(k)} + Y^{(k)}/\mu),\\
S^{(k+1)} &= \mathcal{T}_{\lambda/\mu}(M - L^{(k+1)} + Y^{(k)}/\mu),\\
Y^{(k+1)} &= Y^{(k)} + \mu(M - L^{(k+1)} - S^{(k+1)}),\\
\mu^{(k+1)} &= \min(\rho\,\mu^{(k)},\;\mu_{\max}),
\end{aligned}$$
where $\mathcal{D}_\tau$ soft-thresholds singular values and
$\mathcal{T}_\epsilon(x) = \text{sign}(x)\max(|x|-\epsilon,0)$ is element-wise soft-thresholding.

Stop when $\|M - L^{(k)} - S^{(k)}\|_F / \|M\|_F \le \text{tolerance}$.

The penalty parameter is initialized following Lin et al. (2010):
$$\mu_0 = 1.25 \;/\; \|M\|_2,$$
where $\|M\|_2$ is the spectral norm (largest singular value of $M$).
This ties the initial SVT threshold directly to the scale of the data,
ensuring the low-rank and sparse components are separated correctly from
the very first iteration.

**Step 1: Implement the two thresholding operators.**

- **<font color='red'>singular_value_threshold(A, tau)</font>** — compute SVD of A,
  soft-threshold singular values by tau, return reconstructed matrix.
- **<font color='red'>element_threshold(A, eps)</font>** — element-wise soft-threshold:
  $\text{sign}(A)\max(|A|-\varepsilon,\,0)$.

<!-- grade_id: p2-code-thresh, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: p2-grade-thresh, points: 5, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

# element_threshold
assert_almost_equal(element_threshold(np.array([[0.3,-0.2],[1.0,-1.5]]),0.5),
                    np.array([[0.0,0.0],[0.5,-1.0]]))
# singular_value_threshold: rank should not increase
B = np.outer([1.0,2.0,3.0],[1.0,0.5])  # rank-1
D = singular_value_threshold(B, 0.1)
assert D.shape == B.shape
# large tau => zero matrix
assert_almost_equal(singular_value_threshold(np.eye(3)*0.1, 1.0), np.zeros((3,3)))
print('--- TEST PASSED ---')
```

**Step 2: Implement the IALM Robust PCA solver.**

Write **<font color='red'>robust_pca(M, lam, mu, rho, mu_max, tolerance, max_iter, iter_print)</font>**
with defaults:
- `lam = None` → $1/\sqrt{\max(D,N)}$
- `mu = None`  → $1.25 \;/\; \|M\|_2$ &nbsp; (use `svd(M, compute_uv=False)[0]` for $\|M\|_2$)
- `mu_max = None` → $10^7 \times \mu_0$
- `rho = 1.5`, `tolerance = 1e-5`, `max_iter = 500`, `iter_print = 50`

Return **(L, S)**.

<!-- grade_id: p2-code-rpca, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE
raise NotImplementedError()


```

The cell below tests your robust_pca implementation on a small synthetic matrix of size 20 × 25. It constructs a rank-2 ground-truth low-rank component L and a sparse corruption S with 10% non-zero entries set to ±5, then checks that your decomposition recovers L+S=M with relative error below 1%. Complete the code. 

<!-- grade_id: p2-grade-rpca, points: 5, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###

np.random.seed(123)
D_r, N_r = 20, 25
U_r = np.random.randn(D_r, 2); V_r = np.random.randn(2, N_r)
L_true = U_r @ V_r
S_true = np.zeros((D_r, N_r))
idx_c = np.random.choice(D_r*N_r, size=int(0.10*D_r*N_r), replace=False)
S_true.flat[idx_c] = np.random.choice([-5.0, 5.0], size=len(idx_c))
### Fill in two lines of code here to define your robust PCA ###

rel_err = norm(M_t - L_r - S_r, 'fro') / norm(M_t, 'fro')
assert rel_err < 0.01, f'Error too large: {rel_err:.4f}'
print(f'Relative error: {rel_err:.6f}')
print('--- TEST PASSED ---')
```

**Step 3: Denoise a corrupted image.**

Load the astronaut image, convert to grayscale, and corrupt **25%** of pixels by setting
them to 0 or 1 at random. Store the corrupted image in
**<font color='red'>corrupted_gray</font>**.

Apply `robust_pca` (tolerance=1e-4, max_iter=150) and clip the recovered L to [0,1].
Store in **<font color='red'>L_recovered</font>**.

Plot original, corrupted, and recovered side by side.

<!-- grade_id: p2-code-apply, points: 0, grade: false, solution: true -->
```python
# Load and plot the original colorful image 
astro_color = img_as_float(data.astronaut())
# Plot the colorful
plt.imshow(astro_color)
plt.axis('off')
plt.title('Colorful Original')
plt.show()


```

```python
# Convert to grayscale and plot the image
astro_gray  = rgb2gray(astro_color)   # shape (512, 512)

# Plot the grayscale
plt.imshow(astro_gray, cmap='gray')
plt.axis('off')
plt.title('Grayscale Original')
plt.show()
```

```python
# Corrupt 25% of pixels
### YOUR CODE HERE
raise NotImplementedError()
```

```python
# Apply Robust PCA
### YOUR CODE HERE
raise NotImplementedError()


```

<!-- grade_id: p2-grade-apply, points: 3, grade: true, solution: false -->
```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

In the Markdown cell below answer:
1. What does the low-rank component $L$ represent, and what does sparse component $S$ represent?
2. How does Robust PCA differ from standard PCA when outliers are present?
3. What happens qualitatively if $\lambda$ is set much larger or much smaller?

YOUR ANSWER HERE

---

# Part III - Autoencoder for Dimensionality Reduction (30 marks)

\
We implement a **shallow autoencoder** from scratch using only NumPy and apply it to
compress and reconstruct handwritten digit images.

The network has the structure:
$$z = \sigma(W_1 x + b_1), \qquad \hat{x} = W_2 z + b_2,$$
where $\sigma$ is the **sigmoid** activation applied element-wise and the output layer is
**linear**. Parameters: $W_1\in\mathbb{R}^{p\times d}$, $b_1\in\mathbb{R}^p$,
$W_2\in\mathbb{R}^{d\times p}$, $b_2\in\mathbb{R}^d$.

Loss: $\mathcal{L}=\frac{1}{n}\|{\hat{X}}-X\|_F^2$. Optimized via **mini-batch gradient descent**.

---

```python
f = open('digits.npy','rb')
digits_all = np.load(f); f.close()
N_train = 500
X_train = digits_all[:N_train].astype(float).reshape(N_train, -1) 
print('Training data shape:', X_train.shape)
```

**Step 1: Sigmoid and its derivative.**

- **<font color='red'>sigmoid(z)</font>** — element-wise $1/(1+e^{-z})$.
- **<font color='red'>sigmoid_deriv(z)</font>** — element-wise $\sigma(z)(1-\sigma(z))$.

```python
### YOUR CODE HERE
raise NotImplementedError()


```

```python
### DO NOT REMOVE/EDIT THIS CELL ###

z_t = np.array([-2.0, 0.0, 2.0])
assert_almost_equal(sigmoid(z_t), [0.11920292, 0.5, 0.88079708], decimal=6)
assert_almost_equal(sigmoid_deriv(z_t), sigmoid(z_t)*(1-sigmoid(z_t)), decimal=10)
assert_almost_equal(sigmoid_deriv(np.array([0.0])), [0.25], decimal=8)
print('--- TEST PASSED ---')
```

\
**Step 2: Forward pass.**

**<font color='red'>forward(X, W1, b1, W2, b2)</font>** — takes input batch `(n, d)` and returns
**(z, x_hat)**:
- $z = \sigma(X W_1^\top + b_1)$, shape `(n, p)`.
- $\hat{x} = z W_2^\top + b_2$, shape `(n, d)`.

```python
### YOUR CODE HERE
raise NotImplementedError()


```

```python
### DO NOT REMOVE/EDIT THIS CELL ###

np.random.seed(1)
d_t,p_t,n_t = 576,32,10
W1_t=np.random.randn(p_t,d_t)*0.01; b1_t=np.zeros(p_t)
W2_t=np.random.randn(d_t,p_t)*0.01; b2_t=np.zeros(d_t)
X_t=np.random.rand(n_t,d_t)
z_t,xh_t=forward(X_t,W1_t,b1_t,W2_t,b2_t)
assert z_t.shape==(n_t,p_t) and xh_t.shape==(n_t,d_t)
assert z_t.min()>0 and z_t.max()<1
print('--- TEST PASSED ---')
```

\
**Step 3: Backward pass.**

**<font color='red'>backward(X, z, x_hat, W1, b1, W2, b2)</font>** — returns gradients
**(dW1, db1, dW2, db2)** of $\mathcal{L}=\frac{1}{n}\|\hat{X}-X\|_F^2$.

*Hint: start from the output layer error $\delta_2 = \frac{2}{n}(\hat{X}-X)$, then
backpropagate through $W_2$, through the sigmoid, and through $W_1$.*

```python
### YOUR CODE HERE
raise NotImplementedError()

```

**Step 4: Training loop.**

**<font color='red'>train_autoencoder(X, p, lr, n_epochs, batch_size, random_state)</font>**:
- `lr=0.01`, `n_epochs=50`, `batch_size=128`, `random_state=0`
- Initialize weights: Gaussian std=0.01, biases=0.
- Return **(W1, b1, W2, b2, loss_history)**.

After training: plot the loss curve and display **5 original vs 5 reconstructed** digit images
side by side.

```python
### YOUR CODE HERE
raise NotImplementedError()

```

```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

**Step 5: Compare with PCA.**

Using SVD on `X_train`, compute PCA reconstructions with **p = 32** components
(same bottleneck as the autoencoder). Display original, PCA reconstruction, and autoencoder
reconstruction for the same 5 images in three rows.

In the Markdown cell below answer:
1. Which method gives better visual reconstructions? Does this match your expectation?
2. What objective function difference explains the quality gap?
3. A **linear autoencoder** (identity activation instead of sigmoid) is mathematically
   equivalent to which method, and why?

```python
### YOUR CODE HERE
raise NotImplementedError()


```

```python
### DO NOT REMOVE/EDIT THIS CELL ###
```

YOUR ANSWER HERE

```python

```