# MTH793P - Coursework 6-2

This is a template notebook for the computational exercises of Coursework 6, part **2/2** of the module MTH793P, Advanced machine learning. 

In this part we will implement and explore PCA.

**<font color='red'>IMPORTANT NOTE:</font>**<br>
When using **svd** in this part, always pass the parameter: **<font color='red'>full_matrices=False</font>**.
This is important to speed up the computations.


```python
### 
### Required imports
###

import numpy as np
import matplotlib.pyplot as plt
from sklearn.preprocessing import StandardScaler
from sklearn.datasets import load_wine
from numpy.linalg import svd
from scipy import io

from numpy.testing import assert_almost_equal

%matplotlib inline
```

## PCA for dimension reduction
We will first demonstrate PCA on a 13-dimensional dataset, by loading wine dataset from sklearn, see info [here](https://scikit-learn.org/stable/datasets/toy_dataset.html#wine-dataset). <br>
This dataset contains chemical analysis of N=178 different wines by three different cultivators.
<br>
The analysis contains the folowing measurements:
<dd class="field-odd"><ul class="simple">
<li><p>Alcohol</p></li>
<li><p>Malic acid</p></li>
<li><p>Ash</p></li>
<li><p>Alcalinity of ash</p></li>
<li><p>Magnesium</p></li>
<li><p>Total phenols</p></li>
<li><p>Flavanoids</p></li>
<li><p>Nonflavanoid phenols</p></li>
<li><p>Proanthocyanins</p></li>
<li><p>Colour intensity</p></li>
<li><p>Hue</p></li>
<li><p>OD280/OD315 of diluted wines</p></li>
<li><p>Proline</p></li>
</ul>
</dd>
</dl>

So overall, we have N=178 data points, lying in $\mathbb{R}^{D}$, with D=13. We stack all points together into a matrix **<font color='red'>X_wine</font>** $\in \mathbb{R}^{D\times N}$.<br>
We have labels 0,1, or 2 for each wine (clutivator). The true labels are given in **<font color='red'>L_wine</font>**.<br>
We want to see whether PCA can be helpful in the unsupervised task of clustering the 178 wines.

We start by loading the dataset, and printing the first 5 data points, just to get a general impression.

```python
X_wine, L_wine = load_wine(return_X_y=True)
X_wine = X_wine.T
np.set_printoptions(suppress=True)

print('First 5 data points:')
print('--------------------')
print(X_wine[:,0:5])
```

Write a function called **<font color='red'>pc_transform</font>**. The inputs of this functions are:
* **<font color='red'>X</font>** - the data matrix ($D\times N$)
* **<font color='red'>PCS</font>** - the number of principal components we want to use.
The function should find the first **PCS** componenets, and then return the projection of **X** on these principal components.<br>

Note that our assumption in PCA is that the data are centered, so that the average is 0. Make sure you center the data before computing the principal components.

<!-- grade_id: cell-8748ea089f4f12c4, points: 0, grade: false, solution: true -->
```python
def pc_transform(X, PCS):
    ### YOUR CODE HERE
```

The next box contains a sanity check to see that your function works properly.

<!-- grade_id: cell-f69996274b17d125, points: 1, grade: true, solution: false -->
```python
X_test = np.array([[1,4,3,2],[4,6,7,3],[2,0,7,1]])
U_test = pc_transform(X_test, 2)

RES = np.array([[-0.96686536, -1.7687211,   4.94751325, -2.21192679],
 [-1.49231128,  2.5224855,   0.10336553, -1.13353975]])

assert_almost_equal(np.abs(U_test), np.abs(RES))
print("TEST OK!")
```

Compute the projection of **X_wine** along the first 2 principal components.
Place the result in a matrix **<font color='red'>Y</font>** $\in \mathbb{R}^{2\times N}$. <br>
We then plot the projected points, and colour them according to the original labels.<br>

**<font color='blue'>Q:</font>** What do you think, is this projection helpful for clustering?

<!-- grade_id: cell-440b1f14532c5b48, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE

plt.scatter(Y[0,:], Y[1,:], c=L_wine);
```

We will now try to improve the result by normalisation.<br>
The idea is that we take each row of **X_wine**, and standartise it to have mean 0 and variance 1.<br>

**<font color='blue'>Q:</font>** Why might we want to do that? (HINT: take a look at the columns printed above).<br><br>
You man use the class **StandardScaler** from **sklearn.preprocessing** to do so (already imported). <br>
Place the standartised version of **X_wine** in a matrix called **<font color='red'>XS</font>**. Then compute the first 2 principal componets for **XS**, and place the result into **<font color='red'>YS</font>**. <br>
**NOTE:** Make sure you are normalising the **rows** and not the **columns**

**<font color='blue'>Q:</font>** Is this a better representation for the data?

<!-- grade_id: cell-5534d4d6646971c1, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE

plt.scatter(YS[0,:], YS[1,:], c=L_wine);
```